## Week 11 Practice

### Problem 1: Implementing an objective function in NumPy and PyTorch
Implement two Python functions that implements the mathematical functions $f(x) = 3x^\intercal x - x_1 - 4$. One function should work for $x$ a `numpy.ndarray` and the other should work for $x$ a `torch.FloatTensor`.

```python
pip install torch
```

Output:
```
Collecting torch
  Downloading torch-2.5.1-cp312-cp312-win_amd64.whl.metadata (28 kB)
Requirement already satisfied: filelock in c:\users\jajung\anaconda3\lib\site-packages (from torch) (3.13.1)
Requirement already satisfied: typing-extensions>=4.8.0 in c:\users\jajung\anaconda3\lib\site-packages (from torch) (4.11.0)
Requirement already satisfied: networkx in c:\users\jajung\anaconda3\lib\site-packages (from torch) (3.2.1)
Requirement already satisfied: jinja2 in c:\users\jajung\anaconda3\lib\site-packages (from torch) (3.1.4)
Requirement already satisfied: fsspec in c:\users\jajung\anaconda3\lib\site-packages (from torch) (2024.3.1)
Requirement already satisfied: setuptools in c:\users\jajung\anaconda3\lib\site-packages (from torch) (69.5.1)
Collecting sympy==1.13.1 (from torch)
  Downloading sympy-1.13.1-py3-none-any.whl.metadata (12 kB)
Requirement already satisfied: mpmath<1.4,>=1.1.0 in c:\users\jajung\anaconda3\lib\site-packages (from sympy==1.13.1->torch) (1.3.0)
Requirement al
```

<!-- grade_id: cell-182a74ac56520133, points: 0, grade: false, solution: true -->
```python
import torch
import numpy as np


def f_numpy(x):
    # YOUR CODE HERE
    return 3*np.dot(x,x)-x[0]-4


def f_torch(x):
    # YOUR CODE HERE
    return 3* torch.dot(x,x)-x[0]-4

```

<!-- grade_id: cell-e1c4981cbc9d429f, points: 1, grade: true, solution: false -->
```python
assert abs(f_numpy(np.zeros(4)) + 4) < 1e-6
assert abs(f_numpy(np.ones(2)) - 1) < 1e-6

assert abs(f_torch(torch.zeros(4)) + 4) < 1e-6
assert abs(f_torch(torch.ones(2)) - 1) < 1e-6

x0 = (1, -2, 3, 4, 5)
assert abs(f_numpy(np.array(x0)) - f_torch(torch.Tensor(x0))) < 1e-6
```

### Problem 2: Computing gradients manually and automatically
Implement functions to compute the gradient of $f$ at $x$ using numpy and PyTorch. Use autodiff for the latter.

<!-- grade_id: cell-fb660c996d6b60a5, points: 0, grade: false, solution: true -->
```python
def f_grad_numpy(x):
    # YOUR CODE HERE
    gradient = 6 * x
    gradient[0] -=1
    return gradient


def f_grad_torch(x):
    # YOUR CODE HERE
    f_x = 3 * torch.dot(x,x) -x[0]-4
    gradient = torch.autograd.grad(f_x,x)[0]
    print("Gradient", gradient)
    return gradient
```

<!-- grade_id: cell-27382f3ccab6772b, points: 1, grade: true, solution: false -->
```python
def finite_diff(x, v):
    return (f_numpy(x + v) - f_numpy(x - v)) / (2 * np.linalg.norm(v))


x = np.ones(2)
v0 = 1e-4 * np.array([1, 0])
assert abs(finite_diff(x, v0) - f_grad_numpy(x)[0]) < 1e-2
v1 = 1e-4 * np.array([0, 1])
assert abs(finite_diff(x, v1) - f_grad_numpy(x)[1]) < 1e-2

np.random.seed(42)
for i in range(10):
    x2 = np.random.randn(5)
    v2 = np.random.randn(5)
    v2 = v2 / np.linalg.norm(v2)
    observed = finite_diff(x2, 1e-4 * v2)
    derived_vec = f_grad_numpy(x2)
    derived = derived_vec.dot(v2)
    assert abs(observed - derived) < 1e-2

xt = torch.tensor(x, requires_grad=True)
diff = f_grad_torch(xt).numpy() - f_grad_numpy(x)
assert np.linalg.norm(diff) < 1e-6
```

Output:
```
Gradient tensor([5., 6.], dtype=torch.float64)
```

```python

```